---
title: "1.6 Applications"
author: Jim Bang
output: 
  learnr::tutorial:
    progressive: true
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
library(Ryacas)
library(Deriv)
library(Rsolnp)
gradethis_setup()
tutorial_options(exercise.eval = TRUE, exercise.reveal_solution = TRUE)
knitr::opts_chunk$set(echo = FALSE, fig.width = 6, fig.height = 8,
  fig.align='center', out.width = '75%')
par(mar = c(4, 4, 1, 1))
m = 20; pb = 2; pd = 1
ux <- function(x) - x[1] ^ (alpha) * x[2] ^ (1 - alpha)
bc <- function(x) pb * x[1] + pd * x[2]
alpha = 0.5; t = 0.5
initialEquilibrium <- solnp(c(1, 1), ux, ineqfun = bc,
  ineqLB = 0, ineqUB = m, LB = c(0, 0), UB = c(Inf, Inf))
bcTax <- function(x) (pb + t) * x[1] + pd * x[2]
taxEquilibrium <- solnp(c(1, 1), ux, ineqfun = bcTax, 
  ineqLB = 0, ineqUB = m, LB = c(0, 0), UB = c(Inf, Inf))
alpha = 0.9; sd = 5
beerLoverEquilibrium <- solnp(c(1, 1), ux, ineqfun = bc,
  ineqLB = 0, ineqUB = m, LB = c(0, 0), UB = c(Inf, Inf))
inKindEquilibrium <- solnp(c(1, 1), ux, ineqfun = bc,
    ineqLB = 0, ineqUB = m + sd * pd, LB = c(0, 0), UB = c(m/pb, Inf))
m1 = 80; m2 = 20; r = 0.20; delta = 2/3
ux <- function(x) -(log(x[1]) + delta * log(x[2]))
bc <- function(x) (1 + r) * x[1] + x[2]
intertemporalEquilibrium <- solnp(c(1, 1), ux, ineqfun = bc,
  ineqLB = 0, ineqUB = (1+r) * m1 + m2, LB = c(0,0), UB = c(Inf, Inf))
```

## Overview

There are many, many specific applications for consumer choice. Here, we model a few using our utility maximization framework: 

1. Excise Taxes
2. In-Kind Transfers
3. Insurance and Risk
4. Inter-Temporal Choice 
5. Labor Supply

## Alcoholic Beverage (Excise) Tax

Recall the beer and donuts example from our budget constraint section. Consider a consumer with the income ($m = 20$) and price ($p_b = 2$ and $p_d = 1$) parameters from that example along with preferences described by a Cobb-Douglas utility function with equal budget shares on each good. 

$$\max_x \bigg\{ U(x_b,x_d) = x_b^{0.5}x_d^{0.5} \bigg\}$$
$$\text{s.t. } 2x_b + x_d \le 20$$

### Initial Solution

Calculate and graph the initial equilibrium (with beer on the x-axis and donuts on the y-axis as before). 

```{r taxInitialEquilibrium, exercise = TRUE}

```

```{r taxInitialEquilibrium-solution}
ux <- function(x)
  - x[1] ^ (alpha) * x[2] ^ (1 - alpha)
bc <- function(x)
  pb * x[1] + pd * x[2]
alpha = 0.5
m = 20
pb = 2
pd = 1
initialEquilibrium <-
  solnp(
    c(1, 1),
    ux,
    ineqfun = bc,
    ineqLB = 0,
    ineqUB = m,
    LB = c(0, 0),
    UB = c(Inf, Inf)
  )
```

```{r taxInitialEquilibrium-check}
grade_code() 
```

### Alcoholic-Beverage Tax Solution

Suppose the government levies a $0.50/unit tax on beer. Calculate the new equilibrium.  

```{r taxEquilibrium, exercise = TRUE}

```

```{r taxEquilibrium-solution}
t = 0.5
bcTax <- function(x)
  (pb + t) * x[1] + pd * x[2]
taxEquilibrium <-
  solnp(
    c(1, 1),
    ux,
    ineqfun = bcTax,
    ineqLB = 0,
    ineqUB = m,
    LB = c(0, 0),
    UB = c(Inf, Inf)
  )
```

```{r taxEquilibrium-check}
grade_code() 
```

### Alcoholic-Beverage Tax Graph

Graph the equilibrium (indifference curve and budget line) before and after the tax. Place beer on the x-axis (labeled "Beer") and donuts on the y-axis (labeled "Donuts). 

```{r taxGraph, exercise = TRUE}

```

```{r taxGraph-solution}
curve(
  (-tail(initialEquilibrium$values, 1)) ^ (1 / (1 - alpha)) / (x ^ (alpha / (1 - alpha))),
  from = 0,
  to = m / pb,
  xlim = c(0, m / pb),
  ylim = c(0, m / pd),
  xlab = 'Beer',
  ylab = 'Donuts',
  col = 'dark green'
)
abline(m / pd,-pb / pd, col = 'red')
curve(
  (-tail(taxEquilibrium$value, 1)) ^ (1 / (1 - alpha)) / (x ^ (alpha / (1 - alpha))),
  add = TRUE,
  lty = 'dashed',
  col = 'dark green'
)
abline(m / pd, -(pb + t) / pd,
       lty = 'dashed', col = 'red')
```

```{r taxGraph-check}
grade_code() 
```

Note that the tax reduces beer consumption and consumer utility. However, since the budget share is constant, the tax does *not* affect donut consumption in this example. A consumer with slightly different preferences (like constant elasticity of substitution, shown below for r = 0.5) might respond differently.

Run the following code to simulate the CES case.

```{r cesTax, exercise = TRUE}
ux <- function(x) {
  -(x[1] ^ r + x[2] ^ r) ^ (1 / r)
}
r = 0.5
bc <- function(x) pb * x[1] + pd * x[2]
alpha = 0.5
m = 20
pb = 2
pd = 1
initialEquilibrium <- solnp(c(1, 1), ux, ineqfun = bc, ineqLB = 0, ineqUB = m, LB = c(0, 0), UB = c(Inf, Inf))
t = 0.5
bcTax <- function(x) (pb + t) * x[1] + pd * x[2]
taxEquilibrium <- solnp(c(1, 1), ux, ineqfun = bcTax, ineqLB = 0, ineqUB = m, LB = c(0, 0), UB = c(Inf, Inf))
curve(
  ((-tail(initialEquilibrium$values, 1)) ^ r - x ^ r) ^ (1/r),
  from = 0,
  to = m / pb,
  xlim = c(0, m / pb),
  ylim = c(0, m / pd),
  xlab = 'Beer',
  ylab = 'Donuts',
  col = 'dark green'
)
abline(m / pd,-pb / pd, col = 'red')
points(initialEquilibrium$pars[1],
       initialEquilibrium$pars[2],
       col = 'blue',
       pch = 20)
curve(
  ((-tail(taxEquilibrium$values, 1)) ^ r - x ^ r) ^ (1/r),
  add = TRUE,
  col = 'dark green'
)
abline(m / pd, -(pb + t) / pd,
       lty = 'dashed', col = 'red')
points(taxEquilibrium$pars[1],
       taxEquilibrium$pars[2],
       col = 'blue',
       pch = 20)
```

```{r cesTax-solution}
ux <- function(x) {
  -(x[1] ^ r + x[2] ^ r) ^ (1 / r)
}
r = 0.5
bc <- function(x) pb * x[1] + pd * x[2]
alpha = 0.5
m = 20
pb = 2
pd = 1
initialEquilibrium <- solnp(c(1, 1), ux, ineqfun = bc, ineqLB = 0, ineqUB = m, LB = c(0, 0), UB = c(Inf, Inf))
t = 0.5
bcTax <- function(x) (pb + t) * x[1] + pd * x[2]
taxEquilibrium <- solnp(c(1, 1), ux, ineqfun = bcTax, ineqLB = 0, ineqUB = m, LB = c(0, 0), UB = c(Inf, Inf))
curve(
  ((-tail(initialEquilibrium$values, 1)) ^ r - x ^ r) ^ (1/r),
  from = 0,
  to = m / pb,
  xlim = c(0, m / pb),
  ylim = c(0, m / pd),
  xlab = 'Beer',
  ylab = 'Donuts',
  col = 'dark green'
)
abline(m / pd,-pb / pd, col = 'red')
points(initialEquilibrium$pars[1],
       initialEquilibrium$pars[2],
       col = 'blue',
       pch = 20)
curve(
  ((-tail(taxEquilibrium$values, 1)) ^ r - x ^ r) ^ (1/r),
  add = TRUE,
  col = 'dark green'
)
abline(m / pd, -(pb + t) / pd,
       lty = 'dashed', col = 'red')
points(taxEquilibrium$pars[1],
       taxEquilibrium$pars[2],
       col = 'blue',
       pch = 20)
```

```{r cesTax-check}
grade_code()
```

## Food Stamps (In-Kind Transfers)

Recall the beer and donuts example from our budget constraint section. Consider a consumer with the income ($m = 20$) and price ($p_b = 2$ and $p_d = 1$) parameters from that example along with preferences described by a Cobb-Douglas utility function for a beer-lover who spends 90% of their budget on beer. 

$$\max_x \bigg\{ U(x_b,x_d) = x_b^{0.9}x_d^{0.9} \bigg\}$$
$$\text{s.t. } 2x_b + x_d \le 20$$

```{r, out.width = '300px'}
knitr::include_graphics("images/barney-beer.gif")
```

### Initial Equilibrium Solution

Calculate the initial equilibrium. 

```{r beerLoverEquilibrium, exercise = TRUE}

```

```{r beerLoverEquilibrium-solution}
alpha = 0.9
m = 20
pb = 2
pd = 1
ux <- function(x) {
  -x[1] ^ (alpha) * x[2] ^ (1 - alpha)
}
bc <- function(x) {
  pb * x[1] + pd * x[2]
}
beerLoverEquilibrium <-
  solnp(
    c(1, 1),
    ux,
    ineqfun = bc,
    ineqLB = 0,
    ineqUB = m,
    LB = c(0, 0),
    UB = c(Inf, Inf)
  )
```

```{r beerLoverEquilibrium-check}
grade_code() 
```

### In-Kind Transfer Solution

Suppose the government passes an in-kind transfer of 5 donuts. Calculate and graph the new equilibrium. Include the initial equilibrium (with dashed line type) on your graph.  

```{r inKindEquilibrium, exercise = TRUE, exercise.reveal_solution = TRUE}

```

```{r inKindEquilibrium-solution}
sd = 5
inKindEquilibrium <-
  solnp(
    c(1, 1),
    ux,
    ineqfun = bc,
    ineqLB = 0,
    ineqUB = m + sd * pd,
    LB = c(0, 0),
    UB = c(m/pb, Inf)
  )
```

```{r inKindEquilibrium-check}
grade_code() 
```

### In-Kind Transfer Graph

Graph the equilibrium (indifference curve and budget line) before and after the transfer. Place beer on the x-axis (labeled "Beer") and donuts on the y-axis (labeled "Donuts). 

```{r inKindGraph, exercise = TRUE}
curve(
  (-tail(beerLoverEquilibrium$values, 1)) ^ (1 / (1 - alpha)) / (x ^ (alpha / (1 - alpha))),
  0,
  (m + sd * pd) / pb,
  xlab = 'Beer',
  ylab = 'Donuts',
  xlim = c(0, (m + sd * pd) / pb),
  ylim = c(0, (m + sd * pd) / pd),
  col = 'dark green'
)
abline(m / pd, -pb / pd,
       col = 'red')
curve(
  (-tail(inKindEquilibrium$values, 1)) ^ (1 / (1 - alpha)) / (x ^ (alpha / (1 - alpha))),
  add = TRUE,
  col = 'dark green'
)
curve(
  sd + m / pd - pb * x / pd,
  from = 0,
  to = m / pb,
  add = TRUE,
  col = 'red'
)
segments(
  x0 = m / pb,
  x1 = m / pb,
  y0 = 0,
  y1 = sd, 
  col = 'red'
)
```

```{r inKindGraph-check}
grade_code() 
```

How does the result for this consumer compare to the outcome of a direct income transfer equal to the monetary *value* of the 5 donuts? 

## Intertemporal Choice

### Overview

Intertemporal choice involves choosing between consuming *now* or consuming *later*. In that respect, it resembles the standard consumer between different *goods*, except that the "different goods" differ in terms of the state - the present versus the future - during which the consumer consumes them. This application illustrates the intuition that (1) consumers tend to prefer current consumption over consuming in the future; and (2) variety-loving consumers will choose a "smooth" consumption path that tries to balance current and future consumption. 

### Budget Constraint

The budget constraint requires that the present-value of consumption be less than or equal to the present-value of income (at market interest rates). 

$$c_1 + \frac{c_2}{1+r} \le m_1 + \frac{m_2}{1+r}.$$

In other words, the second-period consumption must be less than or equal to second-period income plus *net* savings from the first period. 

$$c_2 \le m_2 + (1+r)(m_1-c_1).$$

Graph the budget constraint for a consumer who receives an income of 80 in period 1 and 20 in period 2 at a market interest rate of 20%. Position consumption for period 1 on the x-axis. Use the intercepts of the budget line as your x and y limits. 

```{r intertemporalBC, exercise = TRUE}

```

```{r intertemporalBC-solution}
m1 = 80
m2 = 20
r = 0.20
curve(
  m2 + (1 + r) * (m1 - x),
  xlim = c(0, m2 + (1 + r) * m1),
  ylim = c(0, m2 + (1 + r) * m1),
  xlab = expression(x[1]),
  ylab = expression(x[2])
) 
```

```{r intertemporalBC-check}
grade_code()
```

### Preferences

Most intertemporal choice problems model utility as the sum of the (diminishing-returns) utilities from consumption in each time period. At the same time, consumers *subjectively* discount future consumption. For our example, let's assume a log utility function as given by: 

$$ u(c_1, c_2) = \sum_{t=1}^2 \bigg\{\delta^{t-1}u(c_t)\bigg\} = ln(c_1) + \delta ln(c_2)$$

Graph the indifference curves for these intertemporal preferences for total utility levels of 6 and 7. Assume a subjective discount rate of $\beta = 2/3$. Plot the graph over the values of the intercepts of the budget constraint above. 

```{r intertemporalPreferences, exercise = TRUE, exercise.reveal_solution = TRUE}

```

```{r intertemporalPreferences-solution}
delta = 2/3
curve(
  exp((6 - log(x)) / delta),
  from = 0,
  to = m2 + (1 + r) * m1,
  ylim = c(0, m2 + (1 + r) * m1),
  xlab = expression(x[1]),
  ylab = expression(x[2])
)
curve(
  exp((7 - log(x)) / delta),
  add = TRUE,
  xlab = expression(x[1]),
  ylab = expression(x[2])
)
```

```{r intertemporalPreferences-check}
grade_code()
```

### Equilibrium

Calculate the equilibrium numerically in `R`. Graph the indifference curve and budget constraint at the equilibrium. 

```{r intertemporalEquilibrium, exercise = TRUE, exercise.reveal_solution = TRUE}

```

```{r intertemporalEquilibrium-solution}
ux <- function(x) {
  -(log(x[1]) + delta * log(x[2]))
}
bc <- function(x) {
  (1 + r) * x[1] + x[2]
}
intertemporalEquilibrium <-
  solnp(
    c(1, 1),
    ux,
    ineqfun = bc,
    ineqLB = 0,
    ineqUB = (1+r) * m1 + m2,
    LB = c(0,0), 
    UB = c(Inf, Inf)
  )
```

```{r intertemporalEquilibrium-check}
grade_code()
```

### Intertemporal Graph

```{r intertemporalGraph, exercise = TRUE}

```

```{r intertemporalGraph-solution}
curve(
  exp(((-tail(intertemporalEquilibrium$values, 1)) - log(x)) / delta),
  from = 0,
  to = m2 + (1 + r) * m1,
  ylim = c(0, m2 + (1 + r) * m1),
  xlab = expression(x[1]),
  ylab = expression(x[2])
)
curve(m2 + (1 + r) * (m1 - x), add = T) 
```

```{r intertemporalGraph-check}
grade_code()
```

One thing we can see is how much the allocation changes in response to a change in the interest rate *to 200%*. Run the following code to simulate this change.

```{r intertemporalRateHike, exercise = TRUE}
curve(
  exp(((-tail(intertemporalEquilibrium$values, 1)) - log(x)) / delta),
  from = 0,
  to = m2 + (1 + r) * m1,
  ylim = c(0, m2 + (1 + r) * m1),
  xlab = expression(x[1]),
  ylab = expression(x[2])
)
curve(m2 + (1 + r) * (m1 - x), add = T)
r = 2
newRateEquilibrium <-
  solnp(
    c(1, 1),
    ux,
    ineqfun = bc,
    ineqLB = 0,
    ineqUB = (1+r) * m1 + m2,
    LB = c(0,0), 
    UB = c(Inf, Inf)
  )
curve(
  exp(((-tail(newRateEquilibrium$values, 1)) - log(x)) / delta),add = TRUE)
curve(m2 + (1 + r) * (m1 - x), add = TRUE)
```

```{r intertemporalRateHike-solution, exercise = TRUE}
curve(
  exp(((-tail(intertemporalEquilibrium$values, 1)) - log(x)) / delta),
  from = 0,
  to = m2 + (1 + r) * m1,
  ylim = c(0, m2 + (1 + r) * m1),
  xlab = expression(x[1]),
  ylab = expression(x[2])
)
curve(m2 + (1 + r) * (m1 - x), add = T)
r = 2
newRateEquilibrium <-
  solnp(
    c(1, 1),
    ux,
    ineqfun = bc,
    ineqLB = 0,
    ineqUB = (1+r) * m1 + m2,
    LB = c(0,0), 
    UB = c(Inf, Inf)
  )
curve(
  exp(((-tail(newRateEquilibrium$values, 1)) - log(x)) / delta),add = TRUE)
curve(m2 + (1 + r) * (m1 - x), add = TRUE)
```

```{r intertemporalRateHike-check}
grade_code()
```

From this we can see that *current* consumption is not all that sensitive to the interest rate. 

## Insurance and Risk Premia



## Labor Supply

Alternatives for describing choice of labor supply:

1. Choose amounts of consumption (a "good") and *labor* (a "bad"); 
2. Choose amounts of consumption and *leisure* (both "goods").


## Tax Rebates


