---
title: "1.6 Applications"
editor_options: 
  markdown: 
    wrap: 72
output: 
  learnr::tutorial: 
    progressive: true
runtime: shiny_prerendered
author: 
  name: "Created by [Jim Bang](http://www.github.com/bangecon)"
  email: BangJamesT@sau.edu
  affiliation: "[St. Ambrose University](http://www.sau.edu)"
date: "`r format(Sys.Date(), '%d %b %Y')`"
---

<style type="text/css">
h1{font-size: 24pt}
h2{font-size: 20pt}
h3{font-size: 18pt}
h4,h5,h6{font-size: 16pt}
body{font-size: 16pt}
 # doc-metadata {
  margin-bottom: 10px;
  font-size: 0.66em;
}
</style>

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
gradethis_setup()
tutorial_options(exercise.eval = TRUE, exercise.reveal_solution = FALSE)
knitr::opts_chunk$set(echo = FALSE, fig.width = 6, fig.height = 8,
  fig.align='center', out.width = '75%')
```

## Overview

There are many, many specific applications for consumer choice. 
Here, we model a few using our utility maximization framework: 

1. Excise Taxes
2. In-Kind Transfers
3. Insurance and Risk
4. Inter-Temporal Choice 
5. Labor Supply

## Alcoholic Beverage (Excise) Tax

### Initial Solution

Recall the beer and donuts example from our budget constraint section. 
Consider a consumer with the income ($m = 20$) and price ($p_b = 2$ and $p_d = 1$) parameters from that example along with preferences described by a Cobb-Douglas utility function with equal budget shares on each good. 

$$\max_x \bigg\{ U(x_b,x_d) = x_b^{0.5}x_d^{0.5} \bigg\}$$
$$\text{s.t. } 2x_b + x_d \le 20$$

Calculate and graph the initial equilibrium 
(with beer on the x-axis and donuts on the y-axis as before). 

```{r taxInitialEquilibrium, exercise = TRUE}

```

```{r taxInitialEquilibrium-solution}
library(Rsolnp)
ux <- function(x) - x[1] ^ (a) * x[2] ^ (1 - a)
bc <- function(x) pb * x[1] + pd * x[2]
a = 0.5; m = 20; pb = 2; pd = 1; control = list(trace = 0)
eq0 <- solnp(c(1, 1), ux, ineqfun = bc, ineqLB = 0, ineqUB = m,
    LB = c(0, 0), UB = c(Inf, Inf))
```

```{r taxInitialEquilibrium-check}
grade_code() 
```

### Alcoholic-Beverage Tax Solution

Suppose the government levies a $0.50/unit tax on beer. 
Calculate the new equilibrium.  

```{r taxEquilibrium, exercise = TRUE}

```

```{r taxEquilibrium-setup}
a = 0.5; m = 20; pb = 2; pd = 1; t = 0.5; control = list(trace = 0)
ux <- function(x) - x[1] ^ (a) * x[2] ^ (1 - a)
bc <- function(x) pb * x[1] + pd * x[2]
```

```{r taxEquilibrium-solution}
t = 0.5
bcT <- function(x) (pb + t) * x[1] + pd * x[2]
eqT <- solnp(c(1, 1), ux, ineqfun = bcT, ineqLB = 0, ineqUB = m,
    LB = c(0, 0), UB = c(Inf, Inf))
```

```{r taxEquilibrium-check}
grade_code() 
```

### Alcoholic-Beverage Tax Graph

Graph the equilibrium (indifference curve and budget line) before and after the tax. 
Place beer on the x-axis (labeled "Beer") and donuts on the y-axis (labeled "Donuts). 

```{r taxGraph, exercise = TRUE}

```

```{r taxGraph-setup}
a = 0.5; m = 20; pb = 2; pd = 1; t = 0.5; control = list(trace = 0)
bc <- function(x) pb * x[1] + pd * x[2]
ux <- function(x) - x[1] ^ (a) * x[2] ^ (1 - a)
eq0 <- solnp(c(1, 1), ux, ineqfun = bc, ineqLB = 0, ineqUB = m, LB = c(0, 0), 
    UB = c(Inf, Inf))
bcT <- function(x) (pb + t) * x[1] + pd * x[2]
eqT <- solnp(c(1, 1), ux, ineqfun = bcT, ineqLB = 0, ineqUB = m,
    LB = c(0, 0), UB = c(Inf, Inf))
```

```{r taxGraph-solution}
curve(-ux(eq0$pars) ^ (1 / (1 - a)) / (x ^ (a / (1 - a))), from = 0, to = m / pb, 
  ylim = c(0, m / pd), xlab = 'Beer', ylab = 'Donuts', col = 'darkgreen')
curve(m / pd - pb * x / pd, add = TRUE, col = 'red')
curve(-ux(eqT$pars) ^ (1 / (1 - a)) / (x ^ (a / (1 - a))), add = TRUE, 
  lty = 'dashed', col = 'darkgreen')
curve(m / pd - (pb + t) * x / pd, lty = 'dashed', col = 'red')
```

```{r taxGraph-check}
grade_code() 
```

Note that the tax reduces beer consumption and consumer utility. 
However, since the budget share is constant, the tax does *not* affect donut consumption in this example. 
A consumer with slightly different preferences (like constant elasticity of substitution, shown below for r = 0.5) might respond differently.

Run the following code to simulate the CES case.

```{r cesTax, exercise = TRUE}
r = 0.5
ces <- function(x) -(x[1] ^ r + x[2] ^ r) ^ (1 / r)

cesEq0 <- solnp(c(1, 1), ces, ineqfun = bc, ineqLB = 0, ineqUB = m, 
  LB = c(0, 0), UB = c(Inf, Inf))
cesEqT <- solnp(c(1, 1), ces, ineqfun = bcT, ineqLB = 0, ineqUB = m, 
  LB = c(0, 0), UB = c(Inf, Inf))

curve((-ces(cesEq0$pars) ^ r - x ^ r) ^ (1/r), from = 0, to = m / pb, 
  ylim = c(0, m / pd), xlab = 'Beer', ylab = 'Donuts', col = 'darkgreen')
curve(m / pd  - pb * x / pd, add = TRUE, col = 'red')
points(cesEq0$pars[1], cesEq0$pars[2], col = 'blue', pch = 20)

curve((-ces(cesEqT$pars) ^ r - x ^ r) ^ (1/r), add = TRUE, col = 'darkgreen')
curve(m / pd - (pb + t) * x / pd, lty = 'dashed', col = 'red')
points(cesEqT$pars[1], cesEqT$pars[2], col = 'blue', pch = 20)
```

```{r cesTax-setup}
r = 0.5; m = 20; pb = 2; pd = 1; t = 0.5; control = list(trace = 0)
bc <- function(x) pb * x[1] + pd * x[2]
bcT <- function(x) (pb + t) * x[1] + pd * x[2]
```


```{r cesTax-solution}
r = 0.5
ces <- function(x) -(x[1] ^ r + x[2] ^ r) ^ (1 / r)

cesEq0 <- solnp(c(1, 1), ces, ineqfun = bc, ineqLB = 0, ineqUB = m, 
  LB = c(0, 0), UB = c(Inf, Inf))
cesEqT <- solnp(c(1, 1), ces, ineqfun = bcT, ineqLB = 0, ineqUB = m, 
  LB = c(0, 0), UB = c(Inf, Inf))

curve((-ces(cesEq0$pars) ^ r - x ^ r) ^ (1/r), from = 0, to = m / pb, 
  ylim = c(0, m / pd), xlab = 'Beer', ylab = 'Donuts', col = 'darkgreen')
curve(m / pd  - pb * x / pd, add = TRUE, col = 'red')
points(cesEq0$pars[1], cesEq0$pars[2], col = 'blue', pch = 20)

curve((-ces(cesEqT$pars) ^ r - x ^ r) ^ (1/r), add = TRUE, col = 'darkgreen')
curve(m / pd - (pb + t) * x / pd, lty = 'dashed', col = 'red')
points(cesEqT$pars[1], cesEqT$pars[2], col = 'blue', pch = 20)
```

```{r cesTax-check}
grade_code()
```

## Food Stamps (In-Kind Transfers)

### Initial Equilibrium Solution

Recall the beer and donuts example from our budget constraint section. Consider a consumer with the income ($m = 20$) and price ($p_b = 2$ and $p_d = 1$) parameters from that example along with preferences described by a Cobb-Douglas utility function for a beer-lover who spends 90% of their budget on beer. 

$$\max_x \bigg\{ U(x_b,x_d) = x_b^{0.9}x_d^{0.9} \bigg\}$$
$$\text{s.t. } 2x_b + x_d \le 20$$

```{r, out.width = '300px'}
knitr::include_graphics("images/barney-beer.gif")
```

Calculate the initial equilibrium. 

```{r beerLoverEquilibrium, exercise = TRUE}

```

```{r beerLoverEquilibrium-setup}
m = 20; pb = 2; pd = 1; t = 0.5; control = list(trace = 0)
bcT <- function(x) (pb + t) * x[1] + pd * x[2]
ux <- function(x) - x[1] ^ (a) * x[2] ^ (1 - a)
bc <- function(x) pb * x[1] + pd * x[2]
```

```{r beerLoverEquilibrium-solution}
a = 0.9
eqBarney0 <- solnp(c(1, 1), ux, ineqfun = bc, ineqLB = 0, ineqUB = m,
    LB = c(0, 0), UB = c(Inf, Inf))
```

```{r beerLoverEquilibrium-check}
grade_code() 
```

### In-Kind Transfer Solution

Suppose the government passes an in-kind transfer of 5 donuts. Calculate and graph the new equilibrium. Include the initial equilibrium (with dashed line type) on your graph.  

```{r inKindEquilibrium, exercise = TRUE, exercise.reveal_solution = TRUE}

```

```{r inKindEquilibrium-setup}
a = 0.9; m = 20; pb = 2; pd = 1; t = 0.5; control = list(trace = 0)
bcT <- function(x) (pb + t) * x[1] + pd * x[2]
ux <- function(x) - x[1] ^ (a) * x[2] ^ (1 - a)
bc <- function(x) pb * x[1] + pd * x[2]
eqBarney0 <- solnp(c(1, 1), ux, ineqfun = bc, ineqLB = 0, ineqUB = m,
    LB = c(0, 0), UB = c(Inf, Inf))
```

```{r inKindEquilibrium-solution}
sd = 5
eqBarney1 <- solnp(c(1, 1), ux, ineqfun = bc, ineqLB = 0, ineqUB = m + sd*pd,
  LB = c(0, 0), UB = c(m/pb, Inf))
```

```{r inKindEquilibrium-check}
grade_code() 
```

### In-Kind Transfer Graph

Graph the equilibrium (indifference curve and budget line) before and after the transfer. Place beer on the x-axis (labeled "Beer") and donuts on the y-axis (labeled "Donuts). 

```{r inKindGraph, exercise = TRUE}
curve((-ux(eqBarney0$pars)) ^ (1 / (1 - a)) / (x ^ (a / (1 - a))),
  from = 0, to = (m + sd * pd) / pb, ylim = c(0, (m + sd * pd) / pd),
  xlab = 'Beer', ylab = 'Donuts', col = 'darkgreen')
curve(m / pd - pb * x / pd, col = 'red')
curve((-ux(eqBarney1$pars)) ^ (1 / (1 - a)) / (x ^ (a / (1 - a))),
  add = TRUE, col = 'darkgreen')
curve(sd + m / pd - pb*x / pd, from = 0, to = m / pb, add = TRUE, col = 'red')
segments(x0 = m / pb, x1 = m / pb, y0 = 0, y1 = sd,  col = 'red')
```

```{r inKindGraph-setup}
a = 0.9; m = 20; pb = 2; pd = 1; t = 0.5; control = list(trace = 0)
bcT <- function(x) (pb + t) * x[1] + pd * x[2]
ux <- function(x) - x[1] ^ (a) * x[2] ^ (1 - a)
bc <- function(x) pb * x[1] + pd * x[2]
eqBarney0 <- solnp(c(1, 1), ux, ineqfun = bc, ineqLB = 0, ineqUB = m,
    LB = c(0, 0), UB = c(Inf, Inf))
sd = 5
eqBarney1 <- solnp(c(1, 1), ux, ineqfun = bc, ineqLB = 0, ineqUB = m + sd*pd,
  LB = c(0, 0), UB = c(m/pb, Inf))
```

```{r inKindGraph-solution}
curve((-ux(eqBarney0$pars)) ^ (1 / (1 - a)) / (x ^ (a / (1 - a))),
  from = 0, to = (m + sd * pd) / pb, ylim = c(0, (m + sd * pd) / pd),
  xlab = 'Beer', ylab = 'Donuts', col = 'darkgreen')
curve(m / pd - pb * x / pd, col = 'red')
curve((-ux(eqBarney1$pars)) ^ (1 / (1 - a)) / (x ^ (a / (1 - a))),
  add = TRUE, col = 'darkgreen')
curve(sd + m / pd - pb*x / pd, from = 0, to = m / pb, add = TRUE, col = 'red')
segments(x0 = m / pb, x1 = m / pb, y0 = 0, y1 = sd,  col = 'red')
```

```{r inKindGraph-check}
grade_code() 
```

How does the result for this consumer compare to the outcome of a direct income transfer equal to the monetary *value* of the 5 donuts? 

## Intertemporal Choice

### Overview

Intertemporal choice involves choosing between consuming *now* or consuming *later*. In that respect, it resembles the standard consumer between different *goods*, except that the "different goods" differ in terms of the state - the present versus the future - during which the consumer consumes them. This application illustrates the intuition that (1) consumers tend to prefer current consumption over consuming in the future; and (2) variety-loving consumers will choose a "smooth" consumption path that tries to balance current and future consumption. 

### Budget Constraint

The budget constraint requires that the present-value of consumption be less than or equal to the present-value of income (at market interest rates). 

$$c_1 + \frac{c_2}{1+r} \le m_1 + \frac{m_2}{1+r}.$$

In other words, the second-period consumption must be less than or equal to second-period income plus *net* savings from the first period. 

$$c_2 \le m_2 + (1+r)(m_1-c_1).$$

Graph the budget constraint for a consumer who receives an income of 80 in period 1 and 20 in period 2 at a market interest rate of 20%. Position consumption for period 1 on the x-axis. Use the intercepts of the budget line as your x and y limits. 

```{r intertemporalBC, exercise = TRUE}

```

```{r intertemporalBC-solution}
m1 = 80; m2 = 20; r = 0.20
curve(m2 + (1 + r) * (m1 - x), from = 0, to = m2 + (1 + r) * m1, 
  ylim = c(0, m2 + (1 + r) * m1), xlab = expression(x[1]), ylab = expression(x[2]))
```

```{r intertemporalBC-check}
grade_code()
```

### Preferences

Most intertemporal choice problems model utility as the sum of the (diminishing-returns) utilities from consumption in each time period. At the same time, consumers *subjectively* discount future consumption. For our example, let's assume a log utility function as given by: 

$$ u(c_1, c_2) = \sum_{t=1}^2 \bigg\{\delta^{t-1}u(c_t)\bigg\} = ln(c_1) + \delta ln(c_2)$$

Graph the indifference curves for these intertemporal preferences for total utility levels of 6 and 7. Assume a subjective discount rate of $\beta = 2/3$. Plot the graph over the values of the intercepts of the budget constraint above. 

```{r intertemporalPreferences, exercise = TRUE, exercise.reveal_solution = TRUE}

```

```{r intertemporalPreferences-setup}
m1 = 80; m2 = 20; r = 0.20
```

```{r intertemporalPreferences-solution}
d = 2/3
curve(exp((6 - log(x)) / d), from = 0, to = m2 + (1 + r) * m1, 
  ylim = c(0, m2 + (1 + r) * m1), xlab = expression(x[1]), ylab = expression(x[2]))
curve(exp((7 - log(x)) / d), add = TRUE)
```

```{r intertemporalPreferences-check}
grade_code()
```

### Equilibrium

Calculate the equilibrium numerically in `R`. Graph the indifference curve and budget constraint at the equilibrium. 

```{r intertemporalEquilibrium, exercise = TRUE, exercise.reveal_solution = TRUE}

```

```{r intertemporalEquilibrium-setup}
m1 = 80; m2 = 20; r = 0.20; d = 2/3
```

```{r intertemporalEquilibrium-solution}
lnU <- function(x) -(log(x[1]) + d * log(x[2]))
bc <- function(x) (1 + r) * x[1] + x[2]
eqTime <- solnp(c(1, 1), lnU, ineqfun = bc, ineqLB = 0, 
  ineqUB = (1+r) * m1 + m2, LB = c(0,0), UB = c(Inf, Inf))
```

```{r intertemporalEquilibrium-check}
grade_code()
```

### Intertemporal Graph

```{r intertemporalGraph, exercise = TRUE}

```

```{r intertemporalGraph-solution}
curve(exp(((-lnU(eqTime$pars)) - log(x)) / d), from = 0, to = m2 + (1 + r) * m1, 
  ylim = c(0, m2 + (1 + r) * m1), xlab = expression(x[1]), ylab = expression(x[2]))
curve(m2 + (1 + r) * (m1 - x), add = T) 
```

```{r intertemporalGraph-check}
grade_code()
```

### Effect of a Rate Hike

One thing we can see is how much the allocation changes in response to a change in the interest rate *to 200%*. Run the following code to simulate this change.

```{r intertemporalRateHike, exercise = TRUE}
r = 2
eqRateHike <- solnp(c(1, 1), lnU, ineqfun = bc, ineqLB = 0,
    ineqUB = (1+r) * m1 + m2, LB = c(0,0), UB = c(Inf, Inf))
```

```{r intertemporalRateHike-setup}
m1 = 80; m2 = 20; r = 0.20; d = 2/3
lnU <- function(x) -(log(x[1]) + d * log(x[2]))
bc <- function(x) (1 + r) * x[1] + x[2]
eqTime <- solnp(c(1, 1), lnU, ineqfun = bc, ineqLB = 0, 
  ineqUB = (1+r) * m1 + m2, LB = c(0,0), UB = c(Inf, Inf))
```

```{r intertemporalRateHike-solution}
r = 2
eqRateHike <- solnp(c(1, 1), lnU, ineqfun = bc, ineqLB = 0, 
  ineqUB = (1+r) * m1 + m2, LB = c(0,0), UB = c(Inf, Inf))
```

```{r intertemporalRateHike-check}
grade_code()
```

### Graphical Effect of a Rate Hike

Graph the effect of the rate hike (including the original equilibrium for a point of reference).

```{r intertemporalRateHikeGraph, exercise = TRUE}

```

```{r intertemporalRateHikeGraph-setup}
m1 = 80; m2 = 20; r = 0.2; d = 2/3
lnU <- function(x) -(log(x[1]) + d * log(x[2]))
bc <- function(x) (1 + r) * x[1] + x[2]
eqTime <- solnp(c(1, 1), lnU, ineqfun = bc, ineqLB = 0, 
  ineqUB = (1+r) * m1 + m2, LB = c(0,0), UB = c(Inf, Inf))
r = 2
eqRateHike <- solnp(c(1, 1), lnU, ineqfun = bc, ineqLB = 0, 
  ineqUB = (1+r) * m1 + m2, LB = c(0,0), UB = c(Inf, Inf))
```

```{r intertemporalRateHikeGraph-solution}
r = 0.2
curve(exp(((-lnU(eqTime$pars)) - log(x)) / d), from = 0, to = m2 + (1 + r) * m1, 
  ylim = c(0, m2 + (1 + r) * m1), xlab = expression(x[1]), ylab = expression(x[2]))
curve(m2 + (1 + r) * (m1 - x), add = TRUE)
r = 2
curve(exp(((-lnU(eqRateHike$pars)) - log(x)) / d), add = TRUE)
curve(m2 + (1 + r) * (m1 - x), add = TRUE)
```

```{r intertemporalRateHikeGraph-check}
grade_code()
```

From this we can see that *current* consumption is not all that sensitive to the interest rate. 

## Insurance and Risk Premia

### Overview

Choice under uncertainty involves maximizing utility over different *states* or outcomes of some random event. A simple way of formulating this problem is to consider two possible states: a bad outcome in which the consumer suffers some loss of assets and occurs with probability $\pi$; and a good outcome in which the consumer keeps their entire asset endowment and occurs with probablity $1-\pi$. 

The consumer starts with an initial endowment of assets, $E$. They may choose to insure as much of that endowment against a fixed loss ($L < E$) as they want at price, or *premium*, of $p$. Without insurance, the consumer only enjoys $E - L$ (with probability $\pi$) if the bad outcome occurs; they enjoy $E$ fully (with probability $1-\pi$) if there is no loss. 

The consumer may "buy" additional units of bad-state consumption by insuring the losses at price, $p$ for each dollar of bad-state consumption they choose to insure. The consumer chooses the amount of insurance, $I$ before observing the uncertain state outcome. 

To make the problem more concrete, use the following parameter values: 
1. $E = \$100,000$
2. $L = \$50,000$
3. $\pi = 0.2$
4. $p = 0.2$

### Budget Constraint

The budget constraint for this consumer works out to be: 

$$pc_b + (1-p)c_g \le E(m)$$
where $E(m)$ is the expected income of the consumer, i.e. $E - pL$.

Plot the budget constraint for the consumer. 

```{r insuranceBC, exercise = TRUE}

```

```{r insuranceBC-solution}
E  = 100000
L  =  50000
pi =  0.2
p  =  0.2
# Entire Budget Constraint
curve((E - pi * L) / (1 - p) - (p / (1 - p)) * x,
      from = 0,
      to = (E - p * L) / p,
      xlab = expression(c[b]),
      ylab = expression(c[g])
)
# Relevant Budget Constraint between No Insurance and Full Insurance
curve((E - pi * L) / (1 - p) - (p / (1 - p)) * x,
      from = E - p * L,
      to = E + p * L,
      xlab = expression(c[b]),
      ylab = expression(c[g])
)
```

### Preferences

Suppose that the consumer has preferences that we can express as the square root of good and bad state consumption levels, and that the overall level of utility is based on the *expected value* of whichever state we observe.

$$E[u(c_b, c_g)] = \pi\sqrt{c_b} + (1-\pi)\sqrt{c_g}$$

Plot the indifference curve for 

```{r insuranceUtility, exercise = TRUE}

```

```{r insuranceUtility-solution}
curve(((sqrt(99500) - pi * sqrt(x)) / (1 - pi)) ^ 2, from = 0, to = 100000)
```

```{r insuranceUtility-check}
grade_code()
```

### Equilibrium

```{r insuranceEquilibrium, exercise = TRUE}

```

```{r insuranceEquilibrium-solution}
ux <- function(x) {
  -(pi*sqrt(x[1]) + (1-pi)*sqrt(x[2]))
}
bc <- function(x) {
  p*x[1] + (1-p)*x[2] 
}
insuranceEquilibrium <- solnp(
  pars = c(E, E), 
  fun = ux, 
  ineqfun = bc,
  ineqLB = 0, 
  ineqUB = E - pi*L, 
  LB = c(E - L, E - p*L)
)
```

```{r insuranceEquilibrium-check}

```

## Labor Supply

Alternatives for describing choice of labor supply:

1. Choose amounts of consumption (a "good") and *labor* (a "bad"); 
2. Choose amounts of consumption and *leisure* (both "goods").


## Tax Rebates


