---
title: "1.3 Utility Functions"
author: Jim Bang
output: 
  learnr::tutorial: 
    progressive: true
runtime: shiny_prerendered
---

<style type="text/css">
h1{font-size: 24pt}
h2{font-size: 20pt}
h3{font-size: 18pt}
h4,h5,h6{font-size: 16pt}
body{font-size: 16pt}
</style>

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
library(Deriv)
library(Ryacas)
gradethis_setup()
tutorial_options(exercise.eval = TRUE, exercise.reveal_solution = TRUE)
knitr::opts_chunk$set(echo = FALSE, fig.align = "center", fig.height = 6, fig.width = 6)
knitr::include_graphics
u <- ysym('x1^0.25 * x2^0.75')
MU1 <- deriv(u, 'x1')
MU2 <- deriv(u, 'x2')
print(c(MU1, MU2))
MRS <- (MU1/MU2) |>
  yac_str() |> 
  Simplify() |>
  ysym()
MRS
```

## Utility

Utility is a numerical representation of a consumer's preference rankings. 

### Utility Functions 

A utility function assigns numerical values to a person's preferences. 

It's often easier to say $U(A) > U(B)$ than $A \succ B$.

### Are the Values Meaningful?

Short answer, no.

Specific utility functions often have certain features that lead to specific features of how consumers respond to price and/or income changes. 

### Cardinal Utility: Jeremy Bentham (1780)

```{r, out.width = '300px'}
knitr::include_graphics("images/JeremyBentham.jpg")
```

### Ordinal Utility: Vilfredo Pareto (1906)

```{r, out.width = '300px'}
knitr::include_graphics("images/VilfredoPareto.jpg")
```

## "Well-Behaved" Utility Functions

### General Properties

One function that satisfies the basic assumptions so far is a simple multiplicative utility function: 

$$U(x_1, x_2) = x_1x_2$$

1. Use the `persp(x, y, z) function in r to graph this utility function for $0 \le x_1 \le 10$ and $0 \le x_2 \le 10$. (Hint: use seq() define $x$ and $y$ as sequences; and define $z$ as the outer product of $x$ and $y$.)

2. Graph the indifference curves for this utility function for $U(*) = 10$ and $U(*) = 25$. (Hint: to graph an indifference curve, you need to fix $U(x_1,x_2) = \bar{U}$ and solve for $x_2 = f(x_1| U = \bar{U})$.)

```{r simpleUtility, exercise=TRUE, exercise.reveal_solution = TRUE}
x = seq(0, 10, 0.1) # Create a vector sequence for x
y = seq(0, 10, 0.1) # Create a vector sequence for y
z = outer(x, y)     # Create the grid of utilitiesas the outer product of x and y

```

```{r simpleUtility-solution}
x = seq(0, 10, 0.1)
y = seq(0, 10, 0.1)
z = outer(x, y)
persp(x,
      y,
      z,
      xlab = expression(x1),
      ylab = expression(x2),
      zlab = expression(U(x1, x2)))
plot(
  x,
  10 / x,
  type = 'l',
  ylim = c(0, 10),
  xlab = expression(x[1]),
  ylab = expression(x[2])
)
curve(25 / x, add = TRUE)
```

```{r simpleUtility-check}
grade_this({if (!inherits(.result, c("list"))) {
    fail("Your class of your answer should be a list")}
  if (.result$y[1] != Inf) { 
    fail("The y-intercept should approach infinity.") }
  pass()
})
```

One potential problem with this multiplicative utility function is that it does not have diminishing returns. 

### Cobb-Douglas

The Cobb-Douglas utility function is very similar to the simple multiplicative utility function, except that it typically does have diminishing returns. The general form of the Cobb-Douglas function is: 

$$U(x_1, x_2) = x_1^\alpha x_2^\beta,\text{ where } 0 < \alpha < 1 \text{ and } 0 < \beta < 1.$$

For any Cobb-Douglas function, we can apply a *monotonic transformation* to it so that $\alpha + \beta = 1$, or 

$$U(x_1, x_2) = x_1^\alpha x_2^{1 - \alpha},\text{ where } 0 < \alpha < 1.$$

### Graphical Representation

1. Use the `persp(x, y, z) function in r to graph this utility function for $0 \le x_1 \le 50$, $0 \le x_2 \le 50$, and $\alpha = 0.25$. (Hint: use seq() define $x$ and $y$ as sequences; and define $z$ as the outer product of $x$ and $y$.)

```{r CobbDouglas, exercise=TRUE, exercise.reveal_solution = TRUE}

```

```{r CobbDouglas-solution}
x = seq(0, 10, 0.1)
y = seq(0, 10, 0.1)
z = outer(x^0.25, y^0.75)
persp(x,
      y,
      z,
      xlab = expression(x1),
      ylab = expression(x2),
      zlab = expression(U(x1, x2)))
```

```{r CobbDouglas-check}
grade_this({
  if (!inherits(.result, c("matrix"))) {
    fail("Your class of your answer should be a matrix")}
  if (dim(.result) != c(4,4)) { 
    fail("The result should have 4 rows and 4 columns") }
  if (.result[1,1] != 0.2) { 
    fail("The first position of your result should be 0.2") }
  pass()
})
```

### Indifference Curve

2. Plot the indifference curves for the Cobb-Douglas function with $\alpha = 0.25$ for $U(*) = 10$ and $U(*) = 25$ and $0 \le x_1 \le 50$ and $0 \le x_2 \le 50$.

```{r cdIndifference, exercise=TRUE, exercise.reveal_solution = TRUE}

```

```{r cdIndifference-solution}
curve((10 / x^0.25)^(1/0.75),
  type = 'l',
  ylim = c(0, 50),
  xlim = c(0, 50),
  xlab = expression(x[1]),
  ylab = expression(x[2])
)
curve((25 / x^0.25)^(1/0.75), add = TRUE)
```

```{r cdIndifference-check}
grade_this({if (!inherits(.result, c("list"))) {
    fail("Your class of your answer should be a list")}
  if (.result$y[1] != Inf) { 
    fail("The y-intercept should approach infinity.") }
  pass()
})
```

### Quasi-Linear

Quasi-linear utility functions take the form, $U(x) = v(x_1) + cx_2$. A couple of examples of this sort of utility funciton include: 
$$U(x) = bx_1^{\alpha} + cx_2,\text{ where }0<\alpha<1,\text{ and}$$
$$U(x) = 4x_1-x_1^2 + 2x_2$$
Quasi-linear utility is sometimes used because it gives simple solutions. 

### Constant Elasticity of Substitution

The constant elasticity of substitution has the convenient property that the percentage change in the *relative* demand for good one ($x_1/x_2$) for a given percentage change in its *relative* price, $p_1/p_2$ (known as the *elasticity of substitution*), is the same everywhere. 
$$U(x) = (x_1^r + x_2^r)^{\frac{1}{r}}, $$ 
where $s=\frac{1}{1-r}$ is the elasticity of substitution.

## MRS, Revisited

### Marginal Utility

The *marginal utility* of each good is the amount that utility changes as a result of a (very small?) change in the amount of that good. Mathematically, for *inifinitesimal* changes, marginal utility is the *derivative* of $u(x)$ with respect to $x_i$, or: 
$$MU_i = \frac{d}{dx_i}u(x_1, x_2, ...,x_n)$$

Use the `deriv` function to obtain the function for the marginal utility of $u(x) = x_1^{0.25}x_2^{0.75}$ with respect to $x_1$ and $x_2$. Name these derivatives `MU1` and `MU2`, respectively.

```{r marginalUtility, exercise = TRUE, exercise.reveal_solution = TRUE}

```

```{r marginalUtility-solution}
u <- ysym('x1^0.25 * x2^0.75')
MU1 <- deriv(u, 'x1')
MU2 <- deriv(u, 'x2')
```

```{r marginalUtility-check}
grade_this({if (!inherits(.result, c("yac_symbol"))) {
    fail("Your class of your answer should be a yac_symbol")}
  if (.result != "y: 0.75*x1^0.25*x2^(-0.25)") { 
    fail("The marginal utility of good 2 should be 0.75*(x1/x2)^0.25.") }
  pass()
})
```

$$`r tex(MU1)`$$
$$`r tex(MU2)`$$

<details style="line-height:105%"><summary>Click here to see the code for converting the output to `LaTeX` equation syntax.</summary>

The `Ryacas` package includes a `tex` function that converts R expressions into LaTeX-friendly equations, as long as the object is already a `yac_symbol` oject.

\$\` r tex(MU2)\`\$

</details>

### Marginal Rate of Substitution

The *marginal rate of substitution* between two goods ($MRS_{ij}$) is the ratio of the marginal utilities, $\frac{MU_i}{MU_j}$. The MRS tells us the consumer's willingness to exchange good 1 for good 2 on the margin in such a way that they would be *indifferent*. Hence, the MRS also represents the slope of the indifference curve through any given values of $x_i$ and $x_j$. 

Use the `Simplify` function from the `Deriv` package to get $MRS_{12} = \frac{MU_1}{MU_2}$ for the previous exercise. Remember that your marginal utility results stored as `yac_symbol` objects, so you need to convert them to strings before applying `Simplify`, and then finally converting it back to a `ysym`. Try using piping to keep it readable! <br>(Note: `Ryacas` has a `simplify` function, but it just doesn't work as well for the functions we're trying to simplify.)

```{r marginalRateOfSubstitution, exercise = TRUE, exercise.reveal_solution = TRUE}

```

```{r marginalRateOfSubstitution-solution}
MRS <- (MU1/MU2) |>
  yac_str() |> 
  Simplify() |>
  ysym()
MRS
```

```{r marginalRateOfSubstitution-check}
grade_this({if (!inherits(.result, c("yac_symbol"))) {
    fail("Your class of your answer should be a yac_symbol")}
  if (.result$yacas_cmd != "x2/(3*x1)") { 
    fail("The MRS should be x2/(3*x1)") }
  pass()
})
```

$$`r tex(MRS)`$$

### General Form 

We can do the same thing with the symbolic parameters to get the general form of the MRS for the Cobb-Douglas utility function. Define `v` as the general form of the utility function from the previous exercise using the parameter name "alpha" instead of the numerical value of 0.25 and repeat the steps to produce the MRS, `MRSv`. 

In addition, we can take derivatives with respect to both goods in one line, which will give you a vector of derivative functions. The fancy math name for this is the *gradient* of $f$ with respect to $x$, denoted $\nabla f(x_1, x_2)$. Try this in your code. You can then get the $MRS$ by dividing the marginal utilities (dividing the first element of the gradient by the second). 

```{r generalMRS, exercise=TRUE, exercise.reveal_solution = TRUE}

```

```{r generalMRS-solution}
v <- ysym('x1^alpha * x2^(1-alpha)')
MV <- deriv(v, c('x1', 'x2'))
MRSv <- (MV[1]/MV[2]) |>
  yac_str() |> 
  Simplify() |>
  ysym()
MRSv
```

```{r generalMRS-check}
grade_this({if (!inherits(.result, c("yac_symbol"))) {
    fail("Your class of your answer should be a yac_symbol")}
  if (.result$yacas_cmd != "(alpha*x2)/(x1*(1-alpha))") { 
    fail("The MRS should be (alpha/(1-alpha))*(x2/x1)") }
  pass()
})
```

```{r}
v <- ysym('x1^alpha * x2^(1-alpha)')
MV1 <- deriv(v, 'x1')
MV2 <- deriv(v, 'x2')
MRSv <- (MV1/MV2) |>
  yac_str() |> 
  Simplify() |>
  ysym()
```

$$`r tex(MRSv)`$$
## Monotonic Transformations

### Rank-Order Equivalence 

Many (in fact, infinitely many) functions can describe the same preferences (and give the same equilibria)!

For a given utility function, any *monotonic transformation* - strictly increasing function of that utility function - gives an identical ranking of bundles. Hence, such a transformation to the utility function changes absolutely nothing about that consumer's choice.

### Example: Cobb-Douglas versus Log-Linear

The function $log_b(x)$ is *monotonic* for and base, $b$ and for all $x > 0$.

```{r, fig.width = 6, fig.height=4, out.width='75%'}
par(mar=c(5,4,1,1)) 
curve(log(x), from = 0, to = 100)
```

<details style="line-height:105%"><summary>Click here to see the code for the graph.<p></p></summary>
`par(mar=c(5,4,1,1)) # set margins (bottom, left, top, right)`
`curve(log(x), from = 0, to = 100)`<br>
<br>
</details>

A useful trick with Cobb-Douglas utility is taking logs: 

$$u(x) = x_1^{\alpha}x_2^{1-\alpha}$$
$$ v(x) = ln[u(x)] = ln(x_1^{\alpha}x_2^{1-\alpha})$$
Using the log-law formulas (see list of useful formulas at the end of this tutorial) we have: 

$$v(x) = ln(x_1^\alpha) + ln(x_2^{1-\alpha})=\alpha ln(x_1) + (1-\alpha)ln(x_2)$$

In other words, the utility function becomes log-linear, that is, linear in the (natural) *log* of $x$. 

## Special Cases

### Perfect Substitutes

For perfect substitutes, the indifference-curve map is linear. 

1. Graph the utility function from 0 to 25 donuts and 0 to 25 beers for a consumer who sees them as perfect substitutes as given by the following utility function: 
$$u(x) = (x_1 + x_2)^{0.5}$$
Plot the indifference curve for u = 5

```{r perfectSubstitutes, exercise = TRUE}
### Define the level of u, and sequances for x and y.

### Define z = f(x,y) and plot x, y, z using persp()


```

```{r perfectSubstitutes-solution}
u = 5
x = seq(0, 25, 0.1)
y = seq(0, 25, 0.1)
z = (outer(x, y, FUN = "+"))^0.5
persp(x,
      y,
      z)
curve(u^2 - x,
  type = 'l',
  ylim = c(0, 25),
  xlim = c(0, 25))
```

```{r perfectSubstitutes-check}
grade_this({if (!inherits(.result, c("list"))) {
    fail("Your class of your answer should be a list")}
  if (.result$y[1] != 25) { 
    fail("The y-intercept should equal 25.") }
  pass()
})
```


### Perfect Complements

For perfect complements, the indifference curves are L-shaped. Run the following code to graph a set of two indifference curves for donuts and beer for a consumer who sees them as perfect complements. 

```{r perfectComplements, exercise = TRUE, exercise.reveal_solution = TRUE}
plot(
  x = c(rep(10, 11)),
  y = c(seq(10, 80, 7)),
  type = 'l',
  xlim = c(0, 80),
  ylim = c(0, 80),
  xlab = 'Donuts',
  ylab = 'Beer',
  col = 'red', 
  main = 'Beer and Donuts are Perfect Complements'
) 
segments(
  x0 = 10, 
  x1 = 80, 
  y0 = 10, 
  y1 = 10, 
  col = 'red'
)
segments(
  x0 = 20, 
  x1 = 20, 
  y0 = 20, 
  y1 = 80, 
  col = 'orange'
)
segments(
  x0 = 20, 
  x1 = 80, 
  y0 = 20, 
  y1 = 20, 
  col = 'orange'
)
segments(
  x0 = 30, 
  x1 = 30, 
  y0 = 30, 
  y1 = 80, 
  col = 'yellow'
)
segments(
  x0 = 30, 
  x1 = 80, 
  y0 = 30, 
  y1 = 30, 
  col = 'yellow'
)
segments(
  x0 = 40, 
  x1 = 40, 
  y0 = 40, 
  y1 = 80, 
  col = 'green'
)
segments(
  x0 = 40, 
  x1 = 80, 
  y0 = 40, 
  y1 = 40, 
  col = 'green'
)
segments(
  x0 = 50, 
  x1 = 50, 
  y0 = 50, 
  y1 = 80, 
  col = 'blue'
)
segments(
  x0 = 50, 
  x1 = 80, 
  y0 = 50, 
  y1 = 50, 
  col = 'blue'
)
segments(
  x0 = 60, 
  x1 = 60, 
  y0 = 60, 
  y1 = 80, 
  col = 'violet'
)
segments(
  x0 = 60, 
  x1 = 80, 
  y0 = 60, 
  y1 = 60, 
  col = 'violet'
)
arrows(
  x0 = 10,
  x1 = 60,
  y0 = 10,
  y1 = 60,
  col = 'black'
)
```

```{r perfectComplements-solution}
plot(
  x = c(rep(10, 11)),
  y = c(seq(10, 80, 7)),
  type = 'l',
  xlim = c(0, 80),
  ylim = c(0, 80),
  xlab = 'Donuts',
  ylab = 'Beer',
  col = 'red', 
  main = 'Beer and Donuts are Perfect Complements'
) 
segments(
  x0 = 10, 
  x1 = 80, 
  y0 = 10, 
  y1 = 10, 
  col = 'red'
)
segments(
  x0 = 20, 
  x1 = 20, 
  y0 = 20, 
  y1 = 80, 
  col = 'orange'
)
segments(
  x0 = 20, 
  x1 = 80, 
  y0 = 20, 
  y1 = 20, 
  col = 'orange'
)
segments(
  x0 = 30, 
  x1 = 30, 
  y0 = 30, 
  y1 = 80, 
  col = 'yellow'
)
segments(
  x0 = 30, 
  x1 = 80, 
  y0 = 30, 
  y1 = 30, 
  col = 'yellow'
)
segments(
  x0 = 40, 
  x1 = 40, 
  y0 = 40, 
  y1 = 80, 
  col = 'green'
)
segments(
  x0 = 40, 
  x1 = 80, 
  y0 = 40, 
  y1 = 40, 
  col = 'green'
)
segments(
  x0 = 50, 
  x1 = 50, 
  y0 = 50, 
  y1 = 80, 
  col = 'blue'
)
segments(
  x0 = 50, 
  x1 = 80, 
  y0 = 50, 
  y1 = 50, 
  col = 'blue'
)
segments(
  x0 = 60, 
  x1 = 60, 
  y0 = 60, 
  y1 = 80, 
  col = 'violet'
)
segments(
  x0 = 60, 
  x1 = 80, 
  y0 = 60, 
  y1 = 60, 
  col = 'violet'
)
arrows(
  x0 = 10,
  x1 = 60,
  y0 = 10,
  y1 = 60,
  col = 'black'
)
```

```{r perfectComplements-check}
grade_code()
```

### Neutral Good

With a neutral good, utility does not depend at all on the quantity of the neutral good. Run the following code to graph a set of indifference curves for the case where one good is "neutral." 

```{r neutralGood}

```

```{r neutralGood-solution}
plot(
  x = c(rep(10, 11)),
  y = c(seq(0, 60, 6)),
  type = 'l',
  xlim = c(0, 60),
  ylim = c(0, 60),
  xlab = 'Donuts',
  ylab = 'Non-Alcoholic Beer',
  col = 'red', 
  main = 'Non-Alcoholic Beer is a Neutral Good'
) 
segments(
  x0 = 20, 
  x1 = 20, 
  y0 = 0, 
  y1 = 60, 
  col = 'orange'
)
segments(
  x0 = 30, 
  x1 = 30, 
  y0 = 0, 
  y1 = 60, 
  col = 'yellow'
)
segments(
  x0 = 40, 
  x1 = 40, 
  y0 = 0, 
  y1 = 60, 
  col = 'green'
)
segments(
  x0 = 50, 
  x1 = 50, 
  y0 = 0, 
  y1 = 60, 
  col = 'blue'
)
segments(
  x0 = 60, 
  x1 = 60, 
  y0 = 0, 
  y1 = 60, 
  col = 'violet'
)
arrows(
  x0 = 10, 
  x1 = 60, 
  y0 = 30, 
  y1 = 30
)
```

```{r neutralGood-check}
grade_code()
```

### Economic Bads

```{r economicBad, fig.height = 5, fig.width = 5}
plot(
  x = c(seq(10, 20, 1)),
  y = c(seq(0, 70, 7)),
  type = 'l',
  xlim = c(0, 70),
  ylim = c(0, 70),
  xlab = 'Donuts',
  ylab = 'Pants',
  col = 'red', 
  main = 'Pants as an Economic Bad'
) 
segments(
  x0 = 20, 
  x1 = 30, 
  y0 = 0, 
  y1 = 70, 
  col = 'orange'
)
segments(
  x0 = 30, 
  x1 = 40, 
  y0 = 0, 
  y1 = 70, 
  col = 'yellow'
)
segments(
  x0 = 40, 
  x1 = 50, 
  y0 = 0, 
  y1 = 70, 
  col = 'green'
)
segments(
  x0 = 50, 
  x1 = 60, 
  y0 = 0, 
  y1 = 70, 
  col = 'blue'
)
segments(
  x0 = 60, 
  x1 = 70, 
  y0 = 0, 
  y1 = 70, 
  col = 'violet'
)
arrows(
  x0 = 15.7, 
  x1 = 64.2, 
  y0 = 40, 
  y1 = 30
)
```

```{r economicBad-solution}
plot(
  x = c(seq(10, 20, 1)),
  y = c(seq(0, 70, 7)),
  type = 'l',
  xlim = c(0, 70),
  ylim = c(0, 70),
  xlab = 'Donuts',
  ylab = 'Pants',
  col = 'red', 
  main = 'Pants as an Economic Bad'
) 
segments(
  x0 = 20, 
  x1 = 30, 
  y0 = 0, 
  y1 = 70, 
  col = 'orange'
)
segments(
  x0 = 30, 
  x1 = 40, 
  y0 = 0, 
  y1 = 70, 
  col = 'yellow'
)
segments(
  x0 = 40, 
  x1 = 50, 
  y0 = 0, 
  y1 = 70, 
  col = 'green'
)
segments(
  x0 = 50, 
  x1 = 60, 
  y0 = 0, 
  y1 = 70, 
  col = 'blue'
)
segments(
  x0 = 60, 
  x1 = 70, 
  y0 = 0, 
  y1 = 70, 
  col = 'violet'
)
arrows(
  x0 = 15.7, 
  x1 = 64.2, 
  y0 = 40, 
  y1 = 30
)
```

```{r economicBad-check}
grade_code()
```

