---
title: "2.3 Cost Curves"
author: Jim Bang
output: 
  learnr::tutorial: 
    progressive: true
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
library(Ryacas)
library(Deriv)
library(nloptr)
gradethis_setup(exercise.reveal_solution = TRUE, exercise.eval = TRUE)
knitr::opts_chunk$set(echo = FALSE, fig.height = 6, fig.width = 6)
alpha = 1/3
y0 = 2500
w = 1
r = 2
costMin <- nloptr(x0 = c(1, 1), lb = c(0, 0),
                 eval_f = function(x) w*x[1] + r*x[2], 
                 eval_g_ineq = function(x) 
                   y0 - 100 * (x[1]^alpha * x[2]^(1-alpha)), 
                 opts = list("algorithm"="NLOPT_LN_COBYLA", "xtol_rel"=1.0e-8))
costSchedule <- data.frame(Quantity = seq(0, 5000, 100),
    Labor = NA, Capital = NA, Total_Cost = NA, Average_Cost = NA)
for (i in 1:nrow(costSchedule)) {
  costMinTemp <- nloptr(x0 = c(1, 1), lb = c(0, 0),
                        eval_f = function(x) w * x[1] + r * x[2],
                        eval_g_ineq = function(x)
                          costSchedule$Quantity[i] - 
                          100 * (x[1] ^ alpha * x[2] ^ (1 - alpha)),
                        opts = list("algorithm" = "NLOPT_LN_COBYLA",
                                    "xtol_rel" = 1.0e-8))
  costSchedule$Labor[i] <- costMinTemp$solution[1]
  costSchedule$Capital[i] <- costMinTemp$solution[2]
  costSchedule$Total_Cost[i] <- costMinTemp$objective
  costSchedule$Average_Cost[i] <-
    costSchedule$Total_Cost[i] / costSchedule$Quantity[i]
}
srCostSchedule <- data.frame(Quantity = seq(0, 5000, 20),
    Labor = NA, Capital = rep(25, 251), Total_Cost = NA, Average_Cost = NA)
for (i in 1:nrow(srCostSchedule)) {
  costMinTemp <- nloptr(x0 = ifelse(i == 1, 1, srCostSchedule$Labor[i-1]),
                        lb = 0, eval_f = function(x) w * x[1] + r * 25,
                    eval_g_ineq = function(x)
                      srCostSchedule$Quantity[i] - 
                      100 * (x ^ alpha * 25 ^ (1 - alpha)),
                    opts = list("algorithm" = "NLOPT_LN_COBYLA",
                                "xtol_rel" = 1.0e-8))
  srCostSchedule$Labor[i] <- costMinTemp$solution[1]
  srCostSchedule$Total_Cost[i] <- costMinTemp$objective
}
srCostSchedule$Average_Cost <- 
  srCostSchedule$Total_Cost / srCostSchedule$Quantity
srCostSchedule$Variable_Cost <- w * srCostSchedule$Labor
srCostSchedule$Average_Variable_Cost <- 
  srCostSchedule$Variable_Cost / srCostSchedule$Quantity
```

## Long-Run Cost Curves

A rational firm starts out by taking a long-run view of their market. To do this, it will determine the different levels and combinations of inputs that minimize its long run costs. 

Then, it must budget its fixed inputs (in our example, capital) based on (1) the quantity it expects to sell; and (2) the least-costly way to produce that quantity with its production methods.  

### Initial Solution

For expected sales of $y = 2500$ units of output and inpput prices $w = 1$ and $w = 2$, the firm in our previous example minimized costs at 25 units of labor and 25 units of capital at a total cost of 75. 

```{r, echo=TRUE}
alpha = 1/3; y0 = 2500; w = 1; r = 2
costMin <- nloptr(x0 = c(1, 1), lb = c(0, 0), 
                  eval_f = function(x) w*x[1] + r*x[2], 
                  eval_g_ineq = function(x) 
                    y0 - 100 * (x[1]^alpha * x[2]^(1-alpha)), 
                  opts = list("algorithm"="NLOPT_LN_COBYLA",
                             "xtol_rel"=1.0e-8) )
```

To find costs over several quantities (i.e. a cost *schedule*), all we need to do is solve the cost-minimization problem for many different levels of output. Computationally, we can do this using a *for-loop*. 

Iterate the cost-minimization problem over any values of output to find the optimal levels of labor, capital, and cost for levels of output from 0 to 5,000 in increments of 100. 

1. Initialize a new `data.frame` object named `costSchedule` with variables `Quantity`, `Labor`, `Capital`, `Total_Cost`, and `Average_Cost`. <br> 
Set the values of `Quantity` equal to the sequence of outputs from 0 to 5,000 by increments of 100, and the other variables equal to `NA`.
2. Begin a `for(...) {...}` loop starting at `i = 1` through the number of rows (`nrow()`) in costSchedule that does the following for each `i`:
a. Compute the solution to the cost-minimization problem for the quantity of output in row `i` of `costSchedule$Quantity` and (temporarily) stores the result as `costMinTemp`. 
b. Assigns the $i^th$ row of the `Labor` variable in `costSchedule` equal to the first value in the `solution` variable in the `costMinTemp` object, and assigns the $i^th$ row of `Capital` equal to the second value of `solution`.
c. Assigns the value of the `objective` in `costMinTemp` to the `Total_Cost` variable, and calculates the average cost as `Total_Cost` divided by `Quantity`.

```{r costSchedule, exercise = TRUE, exercise.reveal_solution = TRUE}

```

```{r costSchedule-solution}
costSchedule <-
  data.frame(
    Quantity = seq(from = 0, to = 5000, by = 100),
    Labor = NA,
    Capital = NA,
    Total_Cost = NA,
    Average_Cost = NA
  )
for (i in 1:nrow(costSchedule)) {
  costMinTemp <- nloptr(
    x0 = c(1, 1),
    lb = c(0, 0),
    eval_f = function(x)
      w * x[1] + r * x[2],
    eval_g_ineq = function(x)
      costSchedule$Quantity[i] - 100 * (x[1] ^ alpha * x[2] ^ (1 - alpha)),
    opts = list("algorithm" = "NLOPT_LN_COBYLA", "xtol_rel" =
                  1.0e-8)
  )
  costSchedule$Labor[i] <- costMinTemp$solution[1]
  costSchedule$Capital[i] <- costMinTemp$solution[2]
  costSchedule$Total_Cost[i] <- costMinTemp$objective
}
costSchedule$Average_Cost <- costSchedule$Total_Cost / costSchedule$Quantity
```

```{r costSchedule-check}
grade_code()
```

## Graphical Analysis

Plot total and average costs as functions of quantity. Display the plots on the same plot area using `par(mfrow = c(nrow, ncol))`.

```{r costScheduleGraph, exercise = TRUE, exercise.reveal_solution = TRUE}

```

```{r costScheduleGraph-solution}
par(mfrow = c(1, 2))
plot(costSchedule$Quantity, costSchedule$Total_Cost, type = 'l', ylim = c(0,150), xlab = "Quantity", ylab = "Total Cost")
plot(costSchedule$Quantity, costSchedule$Average_Cost, type = 'l', ylim = c(0,0.05), xlab = "Quantity", ylab = "Average Cost")
```

```{r costScheduleGraph-check}
grade_code()
```

### Graphical Example with Diminishing Returns

A quick way to get an increasing cost function is to redefine the production function so that the sum of the Cobb-Douglas exponents is less than one. This creates *decreasing* (instead of constant) returns to scale, and therefore *increasing* average costs. 

Run the following code to illustrate increasing costs. 

```{r increasingCostsGraph, exercise = TRUE}
increasingCostSchedule <-
  data.frame(
    Quantity = seq(0, 5000, 100),
    Labor = NA,
    Capital = NA,
    Total_Cost = NA,
    Average_Cost = NA
  )
for (i in 1:nrow(increasingCostSchedule)) {
  costMinTemp <- nloptr(
    x0 = c(10, 10),
    lb = c(0, 0),
    eval_f = function(x)
      w * x[1] + r * x[2],
    eval_g_ineq = function(x)
      increasingCostSchedule$Quantity[i] -
      100 * (x[1] ^ 0.2 * x[2] ^ (0.6)),
    opts = list("algorithm" = "NLOPT_LN_COBYLA", "xtol_rel" =
                  1.0e-8)
  )
  increasingCostSchedule$Labor[i] <- costMinTemp$solution[1]
  increasingCostSchedule$Capital[i] <- costMinTemp$solution[2]
  increasingCostSchedule$Total_Cost[i] <- costMinTemp$objective
}
increasingCostSchedule$Average_Cost <-
    increasingCostSchedule$Total_Cost /increasingCostSchedule$Quantity
par(mfrow = c(1, 2))
plot(
  increasingCostSchedule$Quantity,
  increasingCostSchedule$Total_Cost,
  type = 'l',
  ylim = c(0, 400),
  xlab = "Quantity",
  ylab = "Total Cost"
)
plot(
  increasingCostSchedule$Quantity,
  increasingCostSchedule$Average_Cost,
  type = 'l',
  ylim = c(0, 0.08),
  xlab = "Quantity",
  ylab = "Average Cost"
)
```


```{r increasingCostsGraph-solution}
increasingCostSchedule <-
  data.frame(
    Quantity = seq(0, 5000, 100),
    Labor = NA,
    Capital = NA,
    Total_Cost = NA,
    Average_Cost = NA
  )
for (i in 1:nrow(increasingCostSchedule)) {
  costMinTemp <- nloptr(
    x0 = c(10, 10),
    lb = c(0, 0),
    eval_f = function(x)
      w * x[1] + r * x[2],
    eval_g_ineq = function(x)
      increasingCostSchedule$Quantity[i] -
      100 * (x[1] ^ 0.2 * x[2] ^ (0.6)),
    opts = list("algorithm" = "NLOPT_LN_COBYLA", "xtol_rel" =
                  1.0e-8)
  )
  increasingCostSchedule$Labor[i] <- costMinTemp$solution[1]
  increasingCostSchedule$Capital[i] <- costMinTemp$solution[2]
  increasingCostSchedule$Total_Cost[i] <- costMinTemp$objective
}
increasingCostSchedule$Average_Cost <-
    increasingCostSchedule$Total_Cost /increasingCostSchedule$Quantity
par(mfrow = c(1, 2))
plot(
  increasingCostSchedule$Quantity,
  increasingCostSchedule$Total_Cost,
  type = 'l',
  ylim = c(0, 400),
  xlab = "Quantity",
  ylab = "Total Cost"
)
plot(
  increasingCostSchedule$Quantity,
  increasingCostSchedule$Average_Cost,
  type = 'l',
  ylim = c(0, 0.08),
  xlab = "Quantity",
  ylab = "Average Cost"
)
```

### Graphical Example with Increasing Returns

Similarly, we can force a decreasing cost function by redefining the exponents in our Cobb-Douglas function so that their sum is greater than one. Such a technology exhibits *increasing* returns to scale and therefore *decreasing* average costs. 

Run the following code to simulate decreasing costs.  

```{r decreasingCostsGraph, exercise = TRUE}
decreasingCostSchedule <-
  data.frame(
    Quantity = seq(0, 5000, 100),
    Labor = NA,
    Capital = NA,
    Total_Cost = NA,
    Average_Cost = NA
  )
for (i in 1:nrow(decreasingCostSchedule)) {
  costMinTemp <- nloptr(
    x0 = c(10, 10),
    lb = c(0, 0),
    eval_f = function(x)
      w * x[1] + r * x[2],
    eval_g_ineq = function(x)
      decreasingCostSchedule$Quantity[i] -
      100 * (x[1] ^ 0.4 * x[2] ^ (0.8)),
    opts = list("algorithm" = "NLOPT_LN_COBYLA", "xtol_rel" =
                  1.0e-8)
  )
  decreasingCostSchedule$Labor[i] <- costMinTemp$solution[1]
  decreasingCostSchedule$Capital[i] <- costMinTemp$solution[2]
  decreasingCostSchedule$Total_Cost[i] <- costMinTemp$objective
}
decreasingCostSchedule$Average_Cost <-
    decreasingCostSchedule$Total_Cost /decreasingCostSchedule$Quantity
par(mfrow = c(1, 2))
plot(
  decreasingCostSchedule$Quantity,
  decreasingCostSchedule$Total_Cost,
  type = 'l',
  ylim = c(0, 80),
  xlab = "Quantity",
  ylab = "Total Cost"
)
plot(
  decreasingCostSchedule$Quantity,
  decreasingCostSchedule$Average_Cost,
  type = 'l',
  ylim = c(0, 0.04),
  xlab = "Quantity",
  ylab = "Average Cost"
)
```

```{r decreasingCostsGraph-solution}
decreasingCostSchedule <-
  data.frame(
    Quantity = seq(0, 5000, 100),
    Labor = NA,
    Capital = NA,
    Total_Cost = NA,
    Average_Cost = NA
  )
for (i in 1:nrow(decreasingCostSchedule)) {
  costMinTemp <- nloptr(
    x0 = c(10, 10),
    lb = c(0, 0),
    eval_f = function(x)
      w * x[1] + r * x[2],
    eval_g_ineq = function(x)
      decreasingCostSchedule$Quantity[i] -
      100 * (x[1] ^ 0.4 * x[2] ^ (0.8)),
    opts = list("algorithm" = "NLOPT_LN_COBYLA", "xtol_rel" =
                  1.0e-8)
  )
  decreasingCostSchedule$Labor[i] <- costMinTemp$solution[1]
  decreasingCostSchedule$Capital[i] <- costMinTemp$solution[2]
  decreasingCostSchedule$Total_Cost[i] <- costMinTemp$objective
}
decreasingCostSchedule$Average_Cost <-
    decreasingCostSchedule$Total_Cost /decreasingCostSchedule$Quantity
par(mfrow = c(1, 2))
plot(
  decreasingCostSchedule$Quantity,
  decreasingCostSchedule$Total_Cost,
  type = 'l',
  ylim = c(0, 80),
  xlab = "Quantity",
  ylab = "Total Cost"
)
plot(
  decreasingCostSchedule$Quantity,
  decreasingCostSchedule$Average_Cost,
  type = 'l',
  ylim = c(0, 0.04),
  xlab = "Quantity",
  ylab = "Average Cost"
)
```

At the end of the day the nature of costs as the firm scales up its inputs is an empirical question. Businesses seeking to optimize input and output decisions hire professional economists to estimate both sales and costs. 

## Short-Run Cost Curves 

In our example, the firm initially anticipated sales of 2500 units. Based on this expectation it budgeted for 25 "units" of capital. Next, the firm just choose inputs (and outputs) based on its short-run costs. 

### Total Cost and Average Cost

With a fixed input (in the short run), a constant returns *to scale* technology becomes a diminishing returns *to the variable input(s)* technology. This introduces increasing costs, as demonstrated by the increasing slope of total cost. 

Run the following code to simulate and graph the short run total and average cost curves for our example. 

```{r srTotalCost, exercise = TRUE, exercise.reveal_solution = TRUE}
srCostSchedule <-
  data.frame(
    Quantity = seq(from = 0, to = 5000, by = 20),
    Labor = NA,
    Capital = 25,
    Total_Cost = NA,
    Average_Cost = NA
  )
for (i in 1:nrow(srCostSchedule)) {
  costMinTemp <- nloptr(
    x0 = c(1),
    lb = c(0),
    eval_f = function(x)
      w * x[1] + r * 25,
    eval_g_ineq = function(x)
      srCostSchedule$Quantity[i] - 100 * (x ^ alpha * 25 ^ (1 - alpha)),
    opts = list("algorithm" = "NLOPT_LN_COBYLA", "xtol_rel" =
                  1.0e-8)
  )
  srCostSchedule$Labor[i] <- costMinTemp$solution[1]
  srCostSchedule$Total_Cost[i] <- costMinTemp$objective
}
srCostSchedule$Average_Cost <- srCostSchedule$Total_Cost / srCostSchedule$Quantity
plot(
  srCostSchedule$Quantity,
  srCostSchedule$Total_Cost,
  type = 'l',
  ylim = c(0, 250),
  xlab = "Quantity",
  ylab = "Total Cost"
)
plot(
  srCostSchedule$Quantity,
  srCostSchedule$Average_Cost,
  type = 'l',
  ylim = c(0, 0.05),
  xlab = "Quantity",
  ylab = "Average Total Cost"
)
```

```{r srTotalCost-solution}
srCostSchedule <-
  data.frame(
    Quantity = seq(from = 0, to = 5000, by = 20),
    Labor = NA,
    Capital = rep(25, 251),
    Total_Cost = NA,
    Average_Cost = NA
  )
for (i in 1:nrow(srCostSchedule)) {
  costMinTemp <- nloptr(
    x0 = c(1),
    lb = c(0),
    eval_f = function(x)
      w * x[1] + r * 25,
    eval_g_ineq = function(x)
      srCostSchedule$Quantity[i] - 100 * (x ^ alpha * 25 ^ (1 - alpha)),
    opts = list("algorithm" = "NLOPT_LN_COBYLA", "xtol_rel" =
                  1.0e-8)
  )
  srCostSchedule$Labor[i] <- costMinTemp$solution[1]
  srCostSchedule$Total_Cost[i] <- costMinTemp$objective
}
srCostSchedule$Average_Cost <- srCostSchedule$Total_Cost / srCostSchedule$Quantity
plot(
  srCostSchedule$Quantity,
  srCostSchedule$Total_Cost,
  type = 'l',
  ylim = c(0, 250),
  xlab = "Quantity",
  ylab = "Total Cost"
)
plot(
  srCostSchedule$Quantity,
  srCostSchedule$Average_Cost,
  type = 'l',
  ylim = c(0, 0.05),
  xlab = "Quantity",
  ylab = "Average Total Cost"
)
```

### Variable Costs

Variable costs are costs associated with variable inputs (in our example labor). They are costs (and input decisions) that the firm can directly control in the short run. 

Run the following code to simulate and graph total and average variable costs for our example. 

```{r variableCost, exercise=TRUE, exercise.reveal_solution=TRUE}
srCostSchedule$Variable_Cost <- w * srCostSchedule$Labor
srCostSchedule$Average_Variable_Cost <-
    srCostSchedule$Variable_Cost / srCostSchedule$Quantity
plot(
  srCostSchedule$Quantity,
  srCostSchedule$Total_Cost,
  type = 'l',
  ylim = c(0, 250),
  xlab = "Quantity",
  ylab = "Total Cost"
)
lines(
  srCostSchedule$Quantity,
  srCostSchedule$Variable_Cost,
  col = 'purple'
)
plot(
  srCostSchedule$Quantity,
  srCostSchedule$Average_Cost,
  type = 'l',
  ylim = c(0, 0.05),
  xlab = "Quantity",
  ylab = "Average Cost"
)
lines(
  srCostSchedule$Quantity,
  srCostSchedule$Average_Variable_Cost,
  col = 'purple'
)
```

```{r variableCost-solution}
srCostSchedule$Variable_Cost <- w * srCostSchedule$Labor
srCostSchedule$Average_Variable_Cost <-
  srCostSchedule$Variable_Cost / srCostSchedule$Quantity
plot(
  srCostSchedule$Quantity,
  srCostSchedule$Total_Cost,
  type = 'l',
  ylim = c(0, 250),
  xlab = "Quantity",
  ylab = "Total Cost"
)
lines(
  srCostSchedule$Quantity,
  srCostSchedule$Variable_Cost,
  type = 'l',
  col = 'purple'
)
plot(
  srCostSchedule$Quantity,
  srCostSchedule$Average_Cost,
  type = 'l',
  ylim = c(0, 0.05),
  xlab = "Quantity",
  ylab = "Average Cost"
)
lines(srCostSchedule$Quantity,
      srCostSchedule$Average_Variable_Cost,
      col = 'purple')
```

### Marginal Costs

Marginal cost measures the rate at which total (or total vairable) cost changes as the firm increases its output by very small amounts. In our numerical simulation, we really only approximate the instantaneous value of the change because our code requires discrete changes (and because we're not deriving the full cost function symbolically). 

To partially compensate for this, we will average the change between the current cost and the *previous* cost with the change between the current cost and the *next* cost. 

By definition, the *derivative* of the cost function  takes a similar approach, but takes the limit as $\Delta Q \to 0$.

Run the following code to add marginal cost to the cost schedule and to the graphs. 

```{r variableCost, exercise=TRUE, exercise.reveal_solution=TRUE}
# Find TC(Q_{i+1}) - TC(Q_{i-1}) / (2 * dQ)
srCostSchedule$Marginal_Cost <- 
  c(NA,diff(srCostSchedule$Total_Cost, lag = 2)/40, NA)
plot(
  srCostSchedule$Quantity,
  srCostSchedule$Total_Cost,
  type = 'l',
  ylim = c(0, 250),
  xlab = "Quantity",
  ylab = "Total Cost"
)
lines(
  srCostSchedule$Quantity,
  srCostSchedule$Variable_Cost,
  type = 'l',
  col = 'purple'
)
plot(
  srCostSchedule$Quantity,
  srCostSchedule$Average_Cost,
  type = 'l',
  ylim = c(0, 0.05),
  xlab = "Quantity",
  ylab = "Average/Marginal Cost"
)
lines(srCostSchedule$Quantity,
      srCostSchedule$Average_Variable_Cost,
      col = 'purple')
lines(srCostSchedule$Quantity,
      srCostSchedule$Marginal_Cost,
      col = 'red')
```

```{r marginalCost-solution}
srCostSchedule$Marginal_Cost <- 
  c(NA,diff(srCostSchedule$Total_Cost, lag = 2)/40, NA)
plot(
  srCostSchedule$Quantity,
  srCostSchedule$Total_Cost,
  type = 'l',
  ylim = c(0, 250),
  xlab = "Quantity",
  ylab = "Total Cost"
)
lines(
  srCostSchedule$Quantity,
  srCostSchedule$Variable_Cost,
  type = 'l',
  col = 'purple'
)
plot(
  srCostSchedule$Quantity,
  srCostSchedule$Average_Cost,
  type = 'l',
  ylim = c(0, 0.05),
  xlab = "Quantity",
  ylab = "Average/Marginal Cost"
)
lines(srCostSchedule$Quantity,
      srCostSchedule$Average_Variable_Cost,
      col = 'purple')
lines(srCostSchedule$Quantity,
      srCostSchedule$Marginal_Cost,
      col = 'red')
```

